# **Funciones**

Este notebook está basado en el curso de [_Julia Academy_](https://github.com/JuliaAcademy/Introduction-to-Julia).

Temas:
1. Como declarar una función en Julia
2. "Duck-typing" en Julia
3. Funciones mutantes vs no mutantes en Julia
4. Algunas funciones de mayor orden

Una función es un bloque de código que podemos llamar para ser ejecutado en nuestro programa y ejecutar una tarea específica. Un ejemplo que hemos utilizado ya por ejemplo es la función `print`, que imprime en pantalla los argumentos que le demos. Aprenderemos ahora a escribir nuestras propias funciones y algunas características útiles que tienen en `Julia`.

## **Como declarar una función en Julia**
Julia tiene varias formas de declarar funciones. La primera es utilizando las palabras clave `function` y `end`.

```julia
function saludo(nombre)
    println("Hola $nombre, ¡qué bueno verte!")
end
```

Output:
```
saludo (generic function with 1 method)
```

`saludo` toma un argumento `nombre`, y llama a la función `print` para escribir un saludo.

```julia
function f(x)
    x^2
end
#Nota: podríamos haber escrito return x^2, y también habría funcionado
#Las funciones de Julia retornan el resultado de la última operación que
#hicieron.
```

Output:
```
f (generic function with 1 method)
```

`f` toma un argumento `x`, y devuelve el resultado de la operación `x^2`.

Podemos llamar a estas funciones de la siguiente manera:

```julia
saludo("Mario")
```

Output:
```
Hola Mario, ¡qué bueno verte!
```

```julia
f(49)
```

Output:
```
2401
```

También podemos definir funciones en una sola línea si es que son suficientemente cortas.

```julia
salu2(nombre) = println("Hola $nombre, ¡qué bueno verte!")
```

Output:
```
salu2 (generic function with 1 method)
```

```julia
f2(x) = x^2
```

Output:
```
f2 (generic function with 1 method)
```

Tanto `salu2` como `f2` son idénticas a `saludo` y `f`. Sólo las hemos definido de una forma más compacta.

```julia
salu2("María")
```

Output:
```
Hola María, ¡qué bueno verte!
```

```julia
f2(42)
```

Output:
```
1764
```

Finalmente, podríamos haberlas declarado como funciones "anónimas"

```julia
saludo3 = nombre -> println("Hola $nombre, ¡qué bueno verte!")
```

Output:
```
#1 (generic function with 1 method)
```

```julia
f3 = x -> x^2
```

Output:
```
#3 (generic function with 1 method)
```

```julia
saludo3("René")
```

Output:
```
Hola René, ¡qué bueno verte!
```

```julia
f3(42)
```

Output:
```
1764
```

El uso principal de las funciones anónimas es para dárselas a otras funciones que tomen funciones como argumento, y veremos ejemplos más adelante. Por el momento no se preocupe mucho por ellas.

Para definir funciones que acepten más argumentos, simplemente debemos escribir los argumentos en el orden que serán aceptados.

```julia
function maximo(x,y)
    return x > y ? x : y
    #===============================================================================
    Nota: aquí hemos utilizado un operador ternario. Esto es equivalente a haber
    escrito
    
    if x > y
        return x
    else
        return y
    end
    ================================================================================#
end
println(maximo("abc","ajk"))
println(maximo(2, 8))
```

Output:
```
ajk
8
```

## **Duck-typing in Julia**
*"Si hace cuak como un pato, es un pato."* <br><br>
Las funciones de Julia van a funcionar siempre que las operaciones estén bien definidas para sus argumentos. <br><br>
Por ejemplo, la función `saludo` que definimos anteriormente funciona si el argumento es un número

```julia
saludo(3.14)
```

Output:
```
Hola 3.14, ¡qué bueno verte!
```

Y `f` funcionará en una matriz. 

```julia
A = rand(3, 3) #Esta línea crea una matriz de 3x3 con elementos aleatorios entre 0 y 1
A
```

Output:
```
3×3 Matrix{Float64}:
 0.63243   0.8037    0.796123
 0.909436  0.252811  0.648774
 0.732501  0.500518  0.684106
```

```julia
f(A)  #f(A) = A^2 = A*A. Notar que la operación que realizó Julia fue la multiplicación
      #de matrices, no la multiplicación elemento por elemento.
```

Output:
```
3×3 Matrix{Float64}:
 1.71404  1.10994  1.56954
 1.2803   1.11955  1.33187
 1.41955  1.05765  1.37588
```

`f` también va a funcionar con un `string`, ya que `^` es el operador utilizado para repetir strings.

```julia
f("hola")  #"hola"^2 = "hola"*"hola" = "holahola"
```

Output:
```
"holahola"
```

Sin embargo, `f` no va a funcionar en un vector. Esto es porque no podemos multiplicar vectores columna por vectores columna.

```julia
v = rand(3)
```

Output:
```
3-element Vector{Float64}:
 0.7158892670845735
 0.06611679545610871
 0.07426176035274323
```

```julia
# Esta línea dará error, ya que v es un vector columna, y v*v no está bien definido.
f(v)
```

Output:
```
@ ./boot.jl:373 [inlined]
 [5] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1196
```

## **Funciones mutantes vs. no-mutantes**

A esta altura puede que haya notado que algunas funciones en los tutoriales anteriores tienen un signo `!`. Por convención, las funciones que terminan en un signo `!` pueden alterar el contenido de sus argumentos, y las que no lo tienen no.

Por ejemplo, miremos la diferencia entre las funciones de ordenamiento `sort` y `sort!`

```julia
v = [3, 5, 2] #Creamos un vector v
```

Output:
```
3-element Vector{Int64}:
 3
 5
 2
```

```julia
sort(v)      #Vemos qué retorna la función sort(v)
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

```julia
v            #Comprobamos que v no fue modificado
```

Output:
```
3-element Vector{Int64}:
 3
 5
 2
```

`sort(v)` retorna un arreglo ordenado de los elementos de `v`, pero `v` no cambia. <br><br>

Por otro lado, si usamos `sort!(v)`, los contenidos de `v` son ordenados dentro del mismo arreglo.

```julia
sort!(v)    #Vemos qué retorna la función sort!(v)
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Vector{Int64}:
 2
 3
 5
```

**Importante**
Es una buena práctica usar esta convención al definir funciones, de manera de saber rápido si la función que está utilizando corre el riesgo de modificar vectores que ha definido. 

## **Comportamiento de las funciones ante arreglos y diccionarios** ###

Veamos qué sucede si intentamos modificar una variable dentro de una función:

```julia
function duplicar(x)
    x = 2*x
end
a = 2
b = duplicar(a)
println("a = $a, b = $b")
```

Output:
```
a = 2, b = 4
```

Notamos que `a` no fue modificada por la función `duplicar`, pese a que internamente realizamos la operación `x = 2*x`. Esto es porque `duplicar` utilizó internamente una copia de `a`, y por lo tanto no modificó la variable original.

Por el contrario, cuando una función recibe como argumento una estructura de datos mutable, puede modificarla. Por ejemplo, la siguiente función duplica los elementos de un arreglo:

```julia
function duplicar!(arreglo)
    for i in 1:size(arreglo,1) #El "1" indica que nos interesa el tamaño en la primer dimensión del arreglo
        arreglo[i] *= 2
        #=======================================================================
        arreglo[i] *= 2 es una forma compacta de escribir
        arreglo[i] = arreglo[i]*2
        Esta notación también funciona para otras operaciones como +, - y /
        =======================================================================#
    end
    return arreglo #Esta línea no es necesaria, solo la agregamos para que
                   #la función retorne algo
end
```

Output:
```
duplicar! (generic function with 1 method)
```

```julia
duplicar!([1,2,3])
```

Output:
```
3-element Vector{Int64}:
 2
 4
 6
```

```julia
w = copy(v)
```

Output:
```
3-element Vector{Int64}:
 56
  4
  6
```

```julia
w[1] = 89
v
```

Output:
```
3-element Vector{Int64}:
 56
  4
  6
```

```julia
v = [1,2,3]
duplicar!(v)
print(v)
```

Output:
```
[2, 4, 6]
```

Como somos buenas personas y no queremos confundir a nadie, agregamos un signo `!`, avisando que nuestra función modificará el argumento de entrada que le demos.
Si no tiene nada mejor que hacer en este momento, pruebe crear una función `duplicar` que retorne el un arreglo multiplicado por 2 sin modificarlo.

También podemos modificar elementos de un diccionario. Por ejemplo, la siguiente función agrega una entrada a un diccionario:

```julia
function agregar_numero!(diccionario, contacto, numero)
    diccionario[contacto] = numero #agregamos un contacto a un diccionario
end
agenda = Dict("Juan" => 5551234, "Martina" => 5558954)
agregar_numero!(agenda, "Sofía", 5559812)
agenda
```

Output:
```
Dict{String, Int64} with 3 entries:
  "Sofía"   => 5559812
  "Juan"    => 5551234
  "Martina" => 5558954
```

Esto es muy útil, ya que Julia no necesita crear copias de los arreglos para trabajar dentro de las funciones, pero es también importante tenerlo en cuenta para no cometer errores y modificar un arreglo cuando no queremos hacerlo.

## **Algunas funciones de mayor orden**

### `map`

`map` es una función de "mayor orden" de Julia que *toma una función* como uno de sus argumentos. `map` luego aplica esa función a todos los elementos de una estructura de datos que le pasemos.
Por ejemplo, ejecutar

```julia
map(f, [1, 2, 3])
```
retornará un arreglo donde `f` fue aplicada a todos los elementos de `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])   #Recordemos que más arriba definimos f(x) = x^2
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

Es decir, elevamos al cuadrado cada elemento de `[1, 2, 3]`.

También podríamos haberle dado a `map` una función anónima,

```julia
map(x -> x^3, [1, 2, 3]) #x -> x^3 es una función anónima, es decir, no tiene nombre.
                         #La definimos únicamente para pasarla como argumento a la
                         #función map
```

Output:
```
3-element Vector{Int64}:
  1
  8
 27
```

y ahora hemos elevado al cubo los elementos de `[1, 2, 3]`

### `broadcast`

`broadcast` es otra función de mayor orden como `map`. `broadcast` es una generaliación de `map`, por lo que puede hacer lo mismo que `map` y mucho más. La sintaxis básica de `broadcast` es la misma que para `map`

```julia
broadcast(f, [1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

Y nuevamente aplicamos `f` a todos los elementos de `[1, 2, 3]`

`broadcast` se diferencia de `map` en cómo se comporta ante multiples argumentos de distintas dimensiones. `broadcast` intentará encontrar una dimensión en común y `map` no. Por ejemplo, notar la diferencia entre estas dos líneas:

```julia
map(+, 1, [1,2,3]) #Nota: Los operadores de Julia son funciones. 
                   #      Escribir +(a,b) es lo mismo que escribir a+b.
```

Output:
```
1-element Vector{Int64}:
 2
```

```julia
broadcast(+,1,[1,2,3])
```

Output:
```
3-element Vector{Int64}:
 2
 3
 4
```

Una abreviatura muy útil de `broadcast` es colocar `.` entre una función a la que quiera hacer `broadcast` y sus argumentos. Por ejemplo,

```julia
broadcast(f, [1, 2, 3])
```
es equivalente a
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Output:
```
3-element Vector{Int64}:
 1
 4
 9
```

Note nuevamente como esto es distinto a llamar
```julia
f([1, 2, 3])
```
Podemos elevar al cuadrado cada elemento del vector, ¡pero esto no es lo mismo que elevar al cuadrado el vector, ya que esto no está definido!

Para mostrar más claramente la diferencia, considere

```julia
f(A)
```
y
```julia
f.(A)
```
para una matriz `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

Output:
```
3×3 Matrix{Int64}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
f(A)
```

Output:
```
3×3 Matrix{Int64}:
  30   36   42
  66   81   96
 102  126  150
```

En este caso estamos haciendo
```
f(A) = A^2 = A * A
``` 

Por otro lado,

```julia
B = f.(A)
```

Output:
```
3×3 Matrix{Int64}:
  1   4   9
 16  25  36
 49  64  81
```

contiene los cuadrados de cada una de las entradas individuales de `A`.

Esta sintaxis nos permite escribir expresiones complejas de una forma mucho más natural

```julia
A .+ 2 .* f.(A) ./ A
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

en lugar de

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

y ambas expresiones van a tener la misma performance.

```julia
@. A + 2 * f(A) / A
```

Output:
```
3×3 Matrix{Float64}:
  3.0   6.0   9.0
 12.0  15.0  18.0
 21.0  24.0  27.0
```

## **Especificando tipos**

Al definir una función, podemos especificar el `tipo` de argumentos que recibe. Incluso, como veremos más adelante en multiple dispatch, podemos hacer que una función se comporte distinto dependiendo del `tipo` de los argumentos que recibe. La sintaxis para especificar que una variable `a` tiene tipo `Type` es `a::Type`.
En el siguiente ejemplo definimos una función que únicamente funciona si el argumento es un número entero

```julia
function factorial(n::Int)
    if n < 1
        println("Error, n < 1. Intente con un número positivo.")
        return -1
    end
    result = 1
    for i in 2:n
        result *= i
    end
    return result
end
```

Output:
```
factorial (generic function with 1 method)
```

```julia
factorial(3)
```

Output:
```
6
```

```julia
factorial(3.0)
```

Output:
```
@ ./boot.jl:373 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1196
```

## **Sobre los argumentos de las funciones**

### **Cantidad arbitraria de variables de argumentos**

La función `print` tiene un comportamiento interesante que quizás haya notado. Podemos darle una cantidad arbitraria de argumentos, y `print` imprime cada uno de ellos:

```julia
print("hola", " yo", " estoy", " dándole", " muchos", " argumentos", " a", " print")
```

Output:
```
hola yo estoy dándole muchos argumentos a print
```

```julia
((1,2,3)...,4,5,6)
```

Output:
```
(1, 2, 3, 4, 5, 6)
```

Esto se logra utilizando el operador _splat_. Aprenderemos de el creando nuestra propia función con una cantidad _variable_ de argumentos. La función `suma_arbitraria` suma una cantidad arbitraria de argumentos:

```julia
function suma_arbitraria(x...)
    println("Argumentos dados: $x")
    return sum(x)
end
```

Output:
```
suma_arbitraria (generic function with 1 method)
```

```julia
suma_arbitraria(1,2,3,4,8,18352)
```

Output:
```
Argumentos dados: (1, 2, 3, 4, 8, 18352)
18370
```

Los elementos de "x" fueron leidos como una tupla. Podríamos tambien haber dado como argumento una tupla o un arreglo:

```julia
suma_arbitraria([i for i in 1:10]...)
```

Output:
```
Argumentos dados: (1, 2, 3, 4, 5, 6, 7, 8, 9, 10)
55
```

Note que si buscamos la documentación de `print`, el segundo argumento de la función es `xs...`.

```julia
?print()
```

Output:
```
[36m  print([io::IO], xs...)[39m

  Write to [36mio[39m (or to the default output stream [36mstdout[39m if [36mio[39m is not given) a
  canonical (un-decorated) text representation. The representation used by
  [36mprint[39m includes minimal formatting and tries to avoid Julia-specific details.

  [36mprint[39m falls back to calling [36mshow[39m, so most types should just define [36mshow[39m.
  Define [36mprint[39m if your type has a separate "plain" representation. For
  example, [36mshow[39m displays strings with quotes, and [36mprint[39m displays strings
  without quotes.

  See also [36mprintln[39m, [36mstring[39m, [36mprintstyled[39m.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> print("Hello World!")[39m
[36m  Hello World![39m
[36m  julia> io = IOBuffer();[39m
[36m  [39m
[36m  julia> print(io, "Hello", ' ', :World!)[39m
[36m  [39m
[36m  julia> String(take!(io))[39m
[36m  "Hello World!"[39m
```

### **Argumentos opcionales**
Podemos dar valores por defecto a ciertas variables, lo cual puede ayudar a quien use la función a evitar dar argumentos innecesarios. La siguiente función retorna la fecha en el formato `(día, mes, año)`, y asume que es enero del año 2000 si no especificamos mes ni año: 

```julia
function fecha(dd, mm = 1, aa = 2000)
    return (dd, mm, aa)
end
```

Output:
```
fecha (generic function with 3 methods)
```

```julia
fecha()    #Esto no funcionará, porque debemos especificar el día
```

Output:
```
@ ./boot.jl:373 [inlined]
 [3] include_string(mapexpr::typeof(REPL.softscope), mod::Module, code::String, filename::String)
   @ Base ./loading.jl:1196
```

```julia
fecha(9) #Aquí, dd = 10, y mm = 1, aa = 2000 por defecto
```

Output:
```
(9, 1, 2000)
```

```julia
fecha(9,3) #Aquí hemos especificado el mes, pero no el año
```

Output:
```
(9, 3, 2000)
```

```julia
fecha(9,3,2022) #Finalmente, aquí especificamos tanto el día como el mes y el año
```

Output:
```
(9, 3, 2022)
```

```julia
fecha((9,3)...)  #Como dato extra, también podemos usar el operador ... para ingresar
                  #una tupla como argumentos separados en cualquier función
```

Output:
```
(9, 3, 2000)
```

### **Argumentos Keyword** (Palabras clave) 

Un problema con la función que definimos anteriormente, es que debemos escribir los argumentos en el orden específico que los espera la función. No solo eso, sino que podemos ahorrarnos escribir el mes y no escribir el año, pero no podemos escribir el año e ignorar el mes. Para solucionar este problema, podemos usar `keyword arguments` tras un signo `;`.

```julia
function fecha_v2(dd; mm = 1, aa = 2000, format = :dma)
    if format == :dma
        return (dd, mm, aa)
    elseif format == :mda
        return (mm, dd, aa)
    elseif format == :amd
        return (aa, mm, dd)
    else
        println("Formato de fecha no reconocido. Solo acepto dma, mda y amd")
    end
end
```

Output:
```
fecha_v2 (generic function with 1 method)
```

```julia
mda = 2
```

Output:
```
2
```

```julia
fecha_v2(1;format = :mda, mm = 3) #El ; solo lo usamos para definir la función, no para llamarla
```

Output:
```
(3, 1, 2000)
```

Ahora podemos escribir solo los argumentos que nos interesen y en el orden que querramos, aunque debemos especificar el nombre del argumento.

```julia
function funcion_kwargs(argumento_1, y...;x...)
    println("el primer argumento vale = $argumento_1")
    println("Usted ingresó los siguientes argumentos keyword")
    
    for element in y
        println(element)
    end
    for (key, val) in x
        println("$key = $val")
    end
end

perro = "perro"

funcion_kwargs(1,2,3,4, a = 3, b = "manzana", perro)

#Note que cualquier cosa luego de un ; será considerado parte de los kwargs!
#La línea anterior sin ; no funciona
```

Output:
```
el primer argumento vale = 1
Usted ingresó los siguientes argumentos keyword
2
3
4
perro
a = 3
b = manzana
```

## **Cómo agregar documentación**
Para documentar una función, simplemente debemos agregar un string antes de la declaración de la función.

```julia
"""
    funciondocumentada(x...)

# Ejemplo de documentación

Esta función existe únicamente para explicar cómo agregar documentación.
Escribir documentación funciona como escribir en modo **markdown**.

Note que para separar líneas debe poner una línea vacía entre medio.

Puede llamarla utilizando `funciondocumentada(argumentos)`.\n

Acepta una cantidad arbitraria de argumentos.
"""
function funciondocumentada(x...)
    pass
end
```

Output:
```
funciondocumentada
```

```julia
@doc funciondocumentada
```

Output:
```
[36m  funciondocumentada(x...)[39m

[1m  Ejemplo de documentación[22m
[1m  ≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡≡[22m

  Esta función existe únicamente para explicar cómo agregar documentación.
  Escribir documentación funciona como escribir en modo [1mmarkdown[22m.

  Note que para separar líneas debe poner una línea vacía entre medio.

  Puede llamarla utilizando [36mfunciondocumentada(argumentos)[39m.

  Acepta una cantidad arbitraria de argumentos.
```

```julia
using JLD2
```

```julia
@doc jldsave
```

Output:
```
[36m  jldsave(filename, compress=false; kwargs...)[39m

  Creates a JLD2 file at [36mfilename[39m and stores the variables given as keyword
  arguments.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  jldsave("example.jld2"; a=1, b=2, c)[39m

  is equivalent to

[36m  jldopen("example.jld2, "w") do f[39m
[36m      f["a"] = 1[39m
[36m      f["b"] = 2[39m
[36m      f["c"] = c[39m
[36m  end[39m

  To choose the io type [36mIOStream[39m instead of the default [36mMmapIO[39m use
  [36mjldsave{IOStream}(fn; kwargs...)[39m.
```

## **Ejercicios opcionales**
***1)*** Escriba una función `suma_uno` que sume 1 a su argumento

```julia

```

```julia
@assert suma_uno(1) == 2
```

```julia
@assert suma_uno(11) == 12
```

***2)*** Use `map` o `broadcast` para incrementar cada elemento de una matriz `A` en `1` y asígnelo a una variable `A1`.

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

***3)***
Use la sintaxis `.` de `broadcast` para incrementar todos los valores de `A1` en `1` y guarde el resultado en `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

***4)*** Cree una función `misuma(vec_1, vec_2)` que retorne la suma directa de `vec_1, vec_2`. Use para esto un loop `for` dentro de la función.
Luego cree otra función llamada `misuma!(vec_r, vec_1, vec_2)` que funcione igual, pero guarde el resultado en `vec_r`.

```julia
#= 
function misuma(vec_1, vec_2)
    *código*
end
=#
```

```julia
#= 
function misuma!(vec_r, vec_1, vec_2)
    *código*
end
=#
```

```julia
#Verificación de misuma. Si esta celda no da error, misuma funciona.
vec_1 = [1,2,3]
vec_2 = [3,2,1]

for elemento in misuma(vec_1, vec_2)
    @assert elemento == 4
end
```

```julia
#Verificación de misuma!. Si esta celda no da error, misuma! funciona.
vec_1 = [1,2,3]
vec_2 = [3,2,1]
vec_r = copy(vec_1)
misuma!(vec_r, vec_1, vec_2)
for elemento in vec_r
    @assert elemento == 4
end
```

***5)*** Repita el ejercicio anterior, pero usando broadcast para escribir una función más compacta

```julia
vec_1 = [1,2,3]
vec_2 = [3,2,1]

for elemento in misuma(vec_1, vec_2)
    @assert elemento == 4
end
```

```julia
vec_1 = [1,2,3]
vec_2 = [3,2,1]
vec_r = copy(vec_1)
misuma!(vec_r, vec_1, vec_2)
for elemento in vec_r
    @assert elemento == 4
end
```

```julia

```

```julia

```